<a href="https://colab.research.google.com/github/CAIKAYUBAKA/hds5210-2023/blob/main/week06/week06_assignment_2.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Week 6 Exercises

_McKinney 6.1_

There are multiple ways to solve the problems below.  You can use any one of several approaches.  For example, you can read CSV files using Pandas or the csv module.  Your score won't depend on which modules you choose to use unless explicitly noted below, but your programming style will still matter.

### 30.1 List of Allergies

In this GitHub repository, there is a file called `allergies.json` that contains a list of patient allergies.  You will need to download this [file from here](https://raw.githubusercontent.com/paulboal/hds5210-2023/main/week06/allergies.json) and then upload it into Google Colab to run these examples. It is taken from sample data provided by the EHR vendor, Epic, here: https://open.epic.com/Clinical/Allergy

Take some time to look at the structure of the file.  You can open it directly in Jupyter by clicking the _Home_ icon, then the _from_instructor_ folder, and then the _data_ folder.

Within the file, you'll see that it is a dictionary with many items in it.  One of those items is called `entry` and that item is a list of things.  You can tell that because the item name is immediately followed by an opening square bracket, signifying the start of a list.  It's line 11 of the file: `  "entry": [`

Write a function named `allergy_count(json_file)` that takes as one parameter the name of the JSON file and returns an integer number of entries in that file.  Your function should open the file, read the json into a Python object, and return how many items there are in the list of `entry`s.

```python
import json
ALLERGIES_FILE="allergies.json"
```

<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
import json

ALLERGIES_FILE = {
  "entry": [
    {
      "resource": {
        "resourceType": "AllergyIntolerance",
        "recordedDate": "2015-08-24T23:11:36Z",
        "status": "confirmed",
        "criticality": "CRITL",
        "onset": "2012-11-07T00:00:00Z",
        "patient": {
          "display": "Jason Argonaut"
        },
        "substance": {
          "text": "PENICILLIN G",
          "coding": [
            {
              "system": "http://www.nlm.nih.gov/research/umls/rxnorm",
              "code": "7980",
              "display": "PENICILLIN G"
            },
            {
              "system": "http://fdasis.nlm.nih.gov",
              "code": "Q42T66VG0C",
              "display": "PENICILLIN G"
            }
          ]
        },
        "reaction": [
          {
            "certainty": "confirmed",
            "onset": "2012-11-07T00:00:00Z",
            "manifestation": [
              {
                "text": "Hives"
              }
            ],
            "note": {
              "text": "Severity low enough to be prescribed if needed."
            }
          }
        ],
        "note": {
          "text": "Severity low enough to be prescribed if needed."
        }
      }
    },
    {
      "resource": {
        "resourceType": "AllergyIntolerance",
        "recordedDate": "2015-11-07T20:55:10Z",
        "status": "confirmed",
        "criticality": "CRITL",
        "id": "TPcWiBG2h2E114Vh0sRT8fQB",
        "onset": "2010-05-02T00:00:00Z",
        "patient": {
          "display": "Jason Argonaut"
        },
        "substance": {
          "text": "SHELLFISH-DERIVED PRODUCTS",
          "coding": [
            {
              "system": "http://hl7.org/fhir/ndfrt",
              "code": "N0000007624",
              "display": "SHELLFISH-DERIVED PRODUCTS"
            }
          ]
        },
        "reaction": [
          {
            "certainty": "confirmed",
            "onset": "2010-05-02T00:00:00Z",
            "manifestation": [
              {
                "text": "Itching"
              }
            ]
          }
        ]
      }
    },
    {
      "resource": {
        "resourceType": "AllergyIntolerance",
        "recordedDate": "2015-11-07T20:56:34Z",
        "status": "confirmed",
        "criticality": "CRITH",
        "id": "TKebKfLXzu6Sp.LY-IpvpmQB",
        "onset": "2014-03-07T00:00:00Z",
        "patient": {
          "display": "Jason Argonaut"
        },
        "substance": {
          "text": "STRAWBERRY",
          "coding": [
            {
              "system": "http://www.nlm.nih.gov/research/umls/rxnorm",
              "code": "892484",
              "display": "STRAWBERRY"
            },
            {
              "system": "http://fdasis.nlm.nih.gov",
              "code": "4J2TY8Y81V",
              "display": "STRAWBERRY"
            }
          ]
        },
        "reaction": [
          {
            "certainty": "confirmed",
            "onset": "2014-03-07T00:00:00Z",
            "manifestation": [
              {
                "text": "Anaphylaxis"
              }
            ]
          }
        ]
      }
    },
    {
      "resource": {
        "resourceType": "AllergyIntolerance",
        "recordedDate": "2015-08-24T23:11:36Z",
        "status": "confirmed",
        "criticality": "CRITL",
        "onset": "1978-12-08T00:00:00Z",
        "patient": {
          "display": "Paul Boal"
        },
        "substance": {
          "text": "PENICILLIN G",
          "coding": [
            {
              "system": "http://www.nlm.nih.gov/research/umls/rxnorm",
              "code": "7980",
              "display": "PENICILLIN G"
            },
            {
              "system": "http://fdasis.nlm.nih.gov",
              "code": "Q42T66VG0C",
              "display": "PENICILLIN G"
            }
          ]
        },
        "reaction": [
          {
            "certainty": "confirmed",
            "onset": "1978-12-08T00:00:00Z",
            "manifestation": [
              {
                "text": "Bruising"
              }
            ],
            "note": {
              "text": "Do not prescribe!"
            }
          }
        ],
        "note": {
          "text": "Do not prescribe!"
        }
      }
    }
  ]
}

def allergy_count(data):
    return len(data['entry'])




```

```python
allergy_count(ALLERGIES_FILE)
```

Output:
```
4
```

<!-- grade_id: q1-tests, points: 5, grade: true, solution: false -->
```python
assert type(allergy_count(ALLERGIES_FILE)) == int
assert allergy_count(ALLERGIES_FILE) == 4
```

### 30.2 Number of Patients

If you dig a little bit deaper into this list of allergies, you'll see that each result has a patient associated with it.  Create a funcation called `patient_count(json_file)` that will count how many unique patients we have in this JSON structure.  

```python
def patient_count(data):
    patients = set()
    for entry in data['entry']:
        patient_name = entry['resource']['patient']['display']
        patients.add(patient_name)
    return len(patients)

```

```python
patient_count(ALLERGIES_FILE)
```

Output:
```
2
```

```python
assert type(patient_count(ALLERGIES_FILE)) == int
assert patient_count(ALLERGIES_FILE) == 2
```

### 30.3 How Many Allergies per Patient

Although each entry is a separate allergy, several of them are for the same patient.  Write a function called `allergy_per_patient(json_file)` that counts up how many allergies each patient has.


```python
def allergy_per_patient(data):
    allergies = {}
    for entry in data['entry']:
        patient_name = entry['resource']['patient']['display']
        if patient_name in allergies:
            allergies[patient_name] += 1
        else:
            allergies[patient_name] = 1
    return allergies

```

```python
allergy_per_patient(ALLERGIES_FILE)
```

Output:
```
{'Jason Argonaut': 3, 'Paul Boal': 1}
```

```python
assert type(allergy_per_patient(ALLERGIES_FILE)) == dict
assert allergy_per_patient(ALLERGIES_FILE) == {'Paul Boal': 1, 'Jason Argonaut': 3}
```

### 30.4 Patient Allergies and Reaction

You'll see in the file that each of the items in the `entry` list have several other attributes including a patient name, substance text representation, and a reaction manifestation.  Create a function named `allergy_list(json_file)` that will create an output list that has patient name, allergy, and reaction for each `entry`.  The actual result you should get will be:

```python
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
```

You'll notice that the reaction and the manifestation of that action are lists.  You only need to capture the first reaction and the first manifestation of the action.  That is, if there is a list of things, just output the first one.

<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
def allergy_list(data):
    allergies_list = []
    for entry in data['entry']:
        patient_name = entry['resource']['patient']['display']
        allergy = entry['resource']['substance']['text']
        reaction = entry['resource']['reaction'][0]['manifestation'][0]['text']
        allergies_list.append([patient_name, allergy, reaction])
    return allergies_list

```

```python
allergy_list(ALLERGIES_FILE)
```

Output:
```
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]
```

<!-- grade_id: q2-tests, points: 10, grade: true, solution: false -->
```python
assert allergy_list(ALLERGIES_FILE) == [['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]

```

### 30.5 Allergy Reaction

Write a function called `allergy_reaction(json_file,patient,substance)` that takes three parameter and returns the reaction that will happen if the patient takes the specified substance.  You can solve this, in part, by calling your `allergy_list` function inside your new `allergy_reaction` function.

If the substance is not found in the allergy list, the function should return None.

<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
def allergy_reaction(json_file, patient, substance):
    allergies = allergy_list(json_file)
    for entry in allergies:
        if entry[0] == patient and entry[1] == substance:
            return entry[2]
    return None

```

```python
allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G')
```

Output:
```
'Hives'
```

<!-- grade_id: q3-tests, points: 5, grade: true, solution: false -->
```python
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G') == 'Hives'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS') == 'Itching'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'STRAWBERRY') == 'Anaphylaxis'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN') == None
assert allergy_reaction(ALLERGIES_FILE, 'Paul Boal', 'PENICILLIN G') == 'Bruising'
```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to save this notebook file back to GitHub.  To do that in Google Colab:
1. File -> Save a Copy in GitHub
2. Make sure your HDS5210 repository is selected
3. Make sure the file name includes the week number like this: `week06/week06_assignment_2.ipynb`
4. Add a commit message that means something

**Be sure week names are lowercase and use a two digit week number!!**

**Be sure you use the same file name provided by the instructor!!**

